# CMPS 140  Yibo Fu 

# Assignment 2

Due Friday May 4th, 11:59pm

## Question 1

![graph](https://docs.google.com/drawings/d/e/2PACX-1vStncj9Nc0LddQeViaYnykNxEZsJoYJMHhub2LLX8s7k7gwYjlnlt0cCcivymFihiZyOOMtHwzk1Z4G/pub?w=960&amp;h=720)

Given the above graph implement the following search algorithms:

1. **Depth First Search** (Graph Style) (Textbook Section 3.4.3)
2. **Iterative Deepening Depth-First Search** (Textbook Section 3.4.5)
3. **A* Search** (A-Star) (Textbook Section 3.5.2)



You will implement each of these search algorithms as a function with the following signature:

```python
def search_algorithm(starting_node, goal_node):
    """
    This function implements a search algorithm.
    Parameters:
    - starting_node: The entry node into the graph
    - goal_node: The integer ID of the goal node.
    
    Returns:
    A list containing the visited node ids in order they were visited with starting node
    always being the first node and the goal node always being the last
    """
    
    visited_nodes_in_order = []
    
    # algorithm implemented here
    
    return visited_nodes_in_order
```


I have put skeleton methods for each of the algorithms for you to fill in.

Here are the pseudo code/descriptions implemenatations from the book.

# Depth First Search
---
![dfs.png](https://lh3.googleusercontent.com/y8b0v4eTG2xfUqkWFm7H_RTb_sJ1ERBiyfWjkClb7GtRUsbAFwpECMnL9JfyTcY4AAhBhCWQiQlLhAmWPQLAN_6Ja_WMFEn5U0PoXaPqNJvKjEf9hW6Meq5xyHdygXqB-WzdBRmVEA=w2400)
![dfs-desc.png](https://lh3.googleusercontent.com/dIDKkPQZ7Vn-HfQ0WPlpthT_RkrQ86lH11hjLceUcDvFF2BCZs2RrfRfEiXgpZVG35vkGkLg6q3nISYeEUriPfwVxSNmhxOO0Gujl43G1lRsN8dTBTaxT8baX8tMu0st49R_CsWBKw=w2400)

# Iterative Deepening Depth First Search
---
![iddfs.png](https://lh3.googleusercontent.com/HrkCyMnoMPZoDELRdqx1cakdgA2rPZ2ImIg5n5XE-A5MXwHch7ZZ6uBE4vmVXv34V6BRFOgyrHl6aTXMm0qijA8s47zda7s9dKDDYzEH1mTnhEPWvOM9r1gE1wwXfXP0_3T22zFO0Q=w2400)
![dls.png](https://lh3.googleusercontent.com/w78lEVXUJitCh8WebB49DACKztyErtXTPi4ybHE0ai5L1orFkdYMtc7AfzdiXF1uuotahgqXoYHeaBpeExztTFR0GvKMIE4ddrYoHe6qSea3zL9TQh6LRYHV69cA3pXGt1nyl3Injg=w2400)

# A* Search
---
![astar](https://lh3.googleusercontent.com/Mp31KIl0a7Eu3ky0WU2QaqMBaAZRFH2ed48ua0dAr-P7GQqwyyYi7MT0NzP-2yaF-YQaykasARoUUUolghYc_Vszd4JCYSbKHQMUaAiT2Np9Khypz7NhwMhDOybJGJPdb1e_6oU1RQ=w2400)

# NOTES
---
 - If two nodes are considered equally good chocices, take the node with the lower ID first. 
 - Do not revisit already explored nodes when implementing the Depth First Search.
 
 

```python
 class Node:
    """
    This class describes a single node contained within a graph. 
    It has the following instannce level attributes:
    
    ID: An integer id for the node i.e. 1
    heuristic_cost: A float value representing the estimated 
                    cost to the goal node
    """    
    def __init__(self, ID, heuristic_cost):
        self.ID = ID
        self.connected_nodes = []
        self.heuristic_cost = heuristic_cost
        
    def __repr__(self):
        ID = self.ID
        hx = self.heuristic_cost
        if len(self.connected_nodes)==0:
            nodes = 'None'
        else:
            nodes = ','.join(str(cn[1].ID) for cn in self.connected_nodes)
        return 'Node:{}\nh(n):{}\nConnected Nodes:{}'.format(ID, hx, nodes)
        
    def set_connected_nodes(self,connected_nodes):
        """
        Adds edges that lead from this node to other nodes:
        
        Parameters:
        - connected_nodes: A list of tuples consisting of (cost, Node), 
                           where 'cost' is a floating point value 
                           indicating the cost to get from this node 
                           to 'Node' and 'Node' is a Node object
        """
        self.connected_nodes = connected_nodes
    
def build_graph():
    """
    Builds the graph to be parsed by the search algorithms.
    Returns: The starting node, which is the entry point into the graph
    """
    ids = range(13)
    coords = [(0,0), (1,1), (1,0), (1,1), (5,2), (3,1), (3,0), 
              (3,-1), (5,1), (4,1), (4,0), (4,-2), (7,0)]
    
    #https://en.wikipedia.org/wiki/Euclidean_distance
    euclidean_distance = lambda x1y1, x2y2: ((x1y1[0]-x2y2[0])**2 +  (x1y1[1]-x2y2[1])**2)**(0.5)
    
    def build_connected_node_list(from_id, to_ids):
        starting_coords = coords[from_id]
        
        connected_nodes = []
        for to_id in to_ids:
            connected_nodes.append((euclidean_distance(starting_coords, coords[to_id]), all_nodes[to_id]))
            
        return connected_nodes
    
    goal_coords = (7,0)
    all_nodes = [Node(_id, euclidean_distance(coord, goal_coords)) for _id, coord in zip(ids, coords)]
    
    all_nodes[8].set_connected_nodes(build_connected_node_list(8, [12]))
    all_nodes[10].set_connected_nodes(build_connected_node_list(10,[12]))
    all_nodes[5].set_connected_nodes(build_connected_node_list(5, [8]))
    all_nodes[6].set_connected_nodes(build_connected_node_list(6, [9, 10]))
    all_nodes[7].set_connected_nodes(build_connected_node_list(7, [11]))
    all_nodes[1].set_connected_nodes(build_connected_node_list(1, [4,5]))
    all_nodes[2].set_connected_nodes(build_connected_node_list(2, [5,6]))
    all_nodes[3].set_connected_nodes(build_connected_node_list(3, [7]))
    all_nodes[0].set_connected_nodes(build_connected_node_list(0, [1,2,3]))
    
    return all_nodes[0]
```

```python
# The starting node. You can use this cell to familiarize
# yourself with the node/graph structure
graph = build_graph()
dir(graph)
graph.ID
```

Output:
```
0
```

<!-- grade_id: cell-8782a75d513b1327, points: 0, grade: false, solution: true -->
```python
def depth_first_search(starting_node, goal_node):
    """
    This function implements the depth first search algorithm
    
    Parameters:
    - starting_node: The entry node into the graph
    - goal_node: The integer ID of the goal node.
    
    Returns:
    A list containing the visited nodes in order they were visited with starting node
    always being the first node and the goal node always being the last
    """
    visited_nodes_in_order = []
    # YOUR CODE HERE
    visited_nodes_in_order.append(starting_node.ID)
    for euclidean_distance, connected_node in starting_node.connected_nodes:
        sub_visited_nodes_in_order = depth_first_search(connected_node, goal_node)
        visited_nodes_in_order.extend(sub_visited_nodes_in_order)
        if goal_node in visited_nodes_in_order:
            break
    return visited_nodes_in_order

def iterative_deepening_depth_first_search(starting_node, goal_node):
    """
    This function implements the iterative deepening depth first search algorithm
    
    Parameters:
    - starting_node: The entry node into the graph
    - goal_node: The integer ID of the goal node.
    
    Returns:
    A list containing the visited node ids in order they were visited with starting node
    always being the first node and the goal node always being the last
    """
    visited_nodes_in_order = []
    # YOUR CODE HERE
    def limited_depth_first_search(starting_node, goal, depth):
        visited_nodes = [starting_node.ID]
        if depth <= 0:
            return visited_nodes
        for euclidean_distance, connected_node in starting_node.connected_nodes:
            sub_visited_nodes = limited_depth_first_search(connected_node, goal, depth-1)
            visited_nodes.extend(sub_visited_nodes)
            if goal in visited_nodes:
                break
        return visited_nodes
    
    depth = 0
    while goal_node not in visited_nodes_in_order:
        sub_visited_nodes_in_order = limited_depth_first_search(starting_node, goal_node, depth)
        visited_nodes_in_order.extend(sub_visited_nodes_in_order)
        depth += 1
    return visited_nodes_in_order

def a_star_search(starting_node, goal_node):
    """
    This function implements the A* search algorithm
    
    Parameters:
    - starting_node: The entry node into the graph
    - goal_node: The integer ID of the goal node.
    
    Returns:
    A list containing the visited node ids in order they were visited with starting node
    always being the first node and the goal nodea always being the last
    """
    visited_nodes_in_order = []
    # YOUR CODE HERE
    open_nodes, close_nodes = [starting_node], []
    f_values = {starting_node: {'g': 0, 'h': 0}}
    parent_node = {}
    while len(open_nodes) > 0:
        current_visit_node = min(open_nodes, key=lambda node:f_values[node]['g'] + f_values[node]['h'])
        visited_nodes_in_order.append(current_visit_node.ID)
        open_nodes.remove(current_visit_node)
        if current_visit_node.ID == goal_node:
            break
        for euclidean_distance, connected_node in current_visit_node.connected_nodes:
            if connected_node not in close_nodes:
                if connected_node in open_nodes:
                    if euclidean_distance < f_values[current_visit_node]['g']:
                        parent_node[connected_node] = current_visit_node
                        f_values[connected_node]['g'] = euclidean_distance
                else:
                    open_nodes.append(connected_node)
                    parent_node[connected_node] = current_visit_node
                    f_values[connected_node] = {'g': euclidean_distance, 'h': connected_node.heuristic_cost}
    #print(visited_nodes_in_order)
    return visited_nodes_in_order
```

<!-- grade_id: cell-0c19db32d84a442e, points: 0, grade: true, solution: false -->
```python
goal_node = 12
depth_first_search_answer = [0, 1, 4, 5, 8, 12]
iterative_deepening_depth_first_search_answer = [0, 0, 1, 2, 3, 0, 1,
                                                 4, 5, 2, 5, 6, 3, 7,
                                                 0, 1, 4, 5, 8, 2, 5,
                                                 8, 6, 9, 10, 3, 7, 11,
                                                 0, 1, 4, 5, 8, 12]
a_star_search_answer = [0, 2, 6, 10, 12]

assert depth_first_search(build_graph(), goal_node)==depth_first_search_answer
assert iterative_deepening_depth_first_search(build_graph(), goal_node)==iterative_deepening_depth_first_search_answer
assert a_star_search(build_graph(), goal_node)==a_star_search_answer
```

# Question 2

Solve the folowing [8-puzzle](https://en.wikipedia.org/wiki/15_puzzle) using steepest ascent hillclimbing. 

![8puzzle.png](https://lh3.googleusercontent.com/cjob9NQOFnv80vfG71y2QBi7blTo9zl_UNcnFz1UhICK59Kc8zWBAfRa26yssTv-b7iQT2TiHlCaOLKPW-fXY_1TdIuT8Yzj9fhpKVwuW9cZSF8ehn0oTY5QA3jUfK-AMhrJu5s2aw=w2400)

1. State the heuristic used.
2. List the steps taken as comma seprated values. For example if the first three moves consist of moving the 3 tile, then the 8 tile, then the 6 tile. Your answer would start like: 3,8,6

Here is the description of the algorithm from the book:

![hill-climbing.png](https://lh3.googleusercontent.com/e1tTwqFAQvLGhfNdq50sQIzx2yxvP0N2Jnxh18qek59-SzrXaO4pUNnANUILUBd5qoPVoEHreBv8eo2f1zMKEtdLqS8xV47dNu-WNpXreqnEGNMCe7mih1AxG4yPdjA_gL6WjWr48w=w2400)

<!-- grade_id: cell-dc1edc8aa056f275, points: 0, grade: true, solution: true -->
[YOUR ANSWER HERE]

### heuistic
sum of all node position to its target position manhattan distance

```python
init_state = [[2, 3, None], [1, 8, 4], [7, 6, 5]]
goal_state = [[1, 2, 3], [8, None, 4], [7, 6, 5]]
init_state_dict = {}
goal_state_dict = {}
for i in range(3):
    for j in range(3):
        init_state_dict[init_state[i][j]] = [i, j]
        goal_state_dict[goal_state[i][j]] = [i, j]
```

```python
def heuistic(state):
    return sum([abs(state[i][0]-goal_state_dict[i][0]) + abs(state[i][1]-goal_state_dict[i][1]) for i in range(1, 9)])

from copy import deepcopy
def get_neighbor_state(state):
    next_states = []
    
    s = [[0, 0, 0], [0, 0, 0], [0, 0, 0]]
    for i in range(1, 9):
        x, y = state[i]
        s[x][y] = i
    x, y = state[None]
    s[x][y] = None
    
    next_x = [0, 2] if x == 1 else [1]
    next_y = [0, 2] if y == 1 else [1]
    for nx in next_x:
        new_state = deepcopy(state)
        new_state[s[nx][y]], new_state[None] = new_state[None], new_state[s[nx][y]]
        next_states.append([new_state, s[nx][y]])
        
    for ny in next_y:
        new_state = deepcopy(state)
        new_state[s[x][ny]], new_state[None] = new_state[None], new_state[s[x][ny]]
        next_states.append([new_state, s[x][ny]])
    
    return next_states
current_state = init_state_dict
```

```python
moves = []
while heuistic(current_state)!=0:
    for next_state, block in get_neighbor_state(current_state):
        if heuistic(next_state) < heuistic(current_state):
            moves.append(block)
            current_state = next_state
            break
",".join([str(x) for x in moves])
```

Output:
```
'3,2,1,8'
```

```python

```